# 1. Método de Gauss-Jordan

Este ejercicio te guiará paso a paso para implementar el método de Gauss-Jordan de diagonalización de matrices.

## Desarrolla en tu ordenador personal

Para el trabajo en grupo usaremos el programa Jupyter notebook.  Puedes instalarlo en tu propio ordenador si lo deseas siguiendo los siguientes pasos:

1. Instala Python 3 del [sitio oficial](http://www.python.org/) (abajo, en la sección *Download* pincha sobre el enlace *Python 3.5.2* o la versión que aparezca junto a *Latest*).  Al instalar aparecerá un cuadro como el que se muestra a continuación.
![Instalación de Python](https://docs.python.org/3/_images/win_installer.png)
Marca la casilla *Add Python 3.5 to PATH* y selecciona la opción *Install Now*.

2. Abre un intérprete de órdenes (*símbolo del sistema*) pulsando la tecla de Windows y tecleando `cmd`.  En el terminal que aparece escribe:
```
pip install jupyter ipywidgets mobilechelonian
exit
```

3. Abre otra vez el intérprete de órdenes (*símbolo del sistema*) y escribe lo siguiente.
```
jupyter nbextension enable --py widgetsnbextension
exit
```

4. Cada vez que quieras usar Jupyter ejecuta el intérprete de órdenes y escribe lo siguiente.
```
cd Documentos
jupyter notebook
```

Esto abrirá automáticamente una pestaña del navegador predeterminado.  Si no lo hace abre el navegador de forma manual y visita la página http://localhost:8888/.

---
## Matriz ampliada

Se pretende resolver el sistema lineal $\mathbf{A}\cdot{}\mathbf{x} = \mathbf{b}$.

Como primer paso calcula la matriz ampliada resultante de la yuxtaposición de la matriz A y la matriz I del mismo tamaño.

<!-- grade_id: imprimir, points: 0, grade: false, solution: true -->
```python
def augm(A):
    """Asume A una lista de listas de números correspondiente a una matriz.
    Imprime la matriz A separando las columnas por un carácter tabulador ('\t').
    
    Por ejemplo: 
    imprimir_matriz([[1,2,3],[1,0,1],[-1,2,0]])
1	2	3
1	0	1
-1	2	0
    """
    ### BEGIN SOLUTION
    assert
    for l in A:
        print('\t'.join([str(i) for i in l]))
    ### END SOLUTION
```

La función debería funcionar correctamente para el ejemplo indicado en la especificación. Compruébalo.

```python
imprimir_matriz([[1,2,3],[1,0,1],[-1,2,0]])
```

Output:
```
1	2	3
1	0	1
-1	2	0
```

<!-- grade_id: imprimir_correcto, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
assert squares(1) == [1]
assert squares(2) == [1, 4]
assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
assert squares(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
```

<!-- grade_id: imprimir_entrada_invalida, points: 1, grade: true, solution: false -->
```python
"""Check that squares raises an error for invalid inputs"""
try:
    imprimir_matriz([[1,2,3],[1,2]])
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    squares(-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

```python
m = [[1,2,3],
    [0,1,1],
    [-1,0,1]]
[len(l) for l in m].count(len(m)) == len(m)
```

Output:
```
True
```

```python
def gauss_jordan(m, epsilon = 1e-10):
    h, w = len(m), len(m[0])
    for y in range(h):
        pivot, maxrow = max([(abs(m[i][y]), i) for i in range(y,h)])
        m[y], m[maxrow] = m[maxrow], m[y]
        assert pivot > epsilon, 'Matriz singular'
        for i in range(y+1, h):
            c = - m[i][y] / m[y][y]
            mac(m[y], c, m[i])

    for y in range(h):
        for i in range(y):
            c = - m[i][y] / m[y][y]
            mac(m[y], c, m[i])
    return m

def mac(v1, k, v2):
    for i in range(len(v1)):
        v2[i] += k*v1[i]
```

```python
gauss_jordan([[0,2,3], [1,-1,3], [1,-1,3]])
```

Output:
```
<ipython-input-66-28b095df473f> in <module>()
----> 1 gauss_jordan([[0,2,3], [1,-1,3], [1,-1,3]])

<ipython-input-64-2f852c3d85c5> in gauss_jordan(m)
     10     for y in range(h):
     11         for i in range(y):
---> 12             c = - m[i][y] / m[y][y]
     13             mac(m[y], c, m[i])
     14     return m

ZeroDivisionError: float division by zero
```

```python
[1,2,3] + [1,1,1]
```

Output:
```
[1, 2, 3, 1, 1, 1]
```

```python

```